<font size="4" style="color:red;"> **IMPORTANT: ** When submitting this homework notebook, please modify only the cells that start with:</font>

```python
# modify this cell
```

# Conditional Probability and Baye's Rule


## Problem

There are two urns $A$ and $B$. Urn $A$ contains $r_A$ red balls and $w_A$ white balls whereas urn $B$ contains $r_B$ red balls and $w_B$ white balls. One of the urns is picked at random and then one ball is picked at random from this urn. Write a function **conditional_probability** that calculates the conditional probability that the randomly chosen ball belonged to urn $A$ given that it is white. Assume that $\frac{r_A}{w_A}\neq\frac{r_B}{w_B}$.

<font  style="color:blue"> **Code**</font>
```python
rA, wA, rB, wB = 1., 2., 2., 1.
conditional__probability(rA, wA, rB, wB) 
```

<font  style="color:magenta"> **Output**</font>
```
0.6666666666666666
```

```python
# modify this cell

def conditional__probability(rA, wA, rB, wB):
    # inputs: all of them are of type 'float'
    # output: a variable of type 'float'
    
    #
    # YOUR CODE HERE
    Omega = rA+wA+rB+wB
    Con_prob = ((rA+wA)/Omega)*(wA/(rA+wA))/((wA+wB)/Omega)
    return Con_prob
    #

```

<!-- grade_id: ex1, points: 5, grade: true, solution: false -->
```python
assert( abs(conditional__probability(2., 4., 3., 3.) -0.5714285714285715) < 10**-5)  
assert( abs(conditional__probability(1., 3., 5., 2.) -0.7241379310344829) < 10**-5) 

#
# AUTOGRADER TEST - DO NOT REMOVE
#

```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-2-608ab43f6d0f> in <module>()
      1 assert( abs(conditional__probability(2., 4., 3., 3.) -0.5714285714285715) < 10**-5)
----> 2 assert( abs(conditional__probability(1., 3., 5., 2.) -0.7241379310344829) < 10**-5)
      3 
      4 #
      5 # AUTOGRADER TEST - DO NOT REMOVE

AssertionError:
```

```python

conditional__probability(2, 1, 6, 4)






```

Output:
```
0.2
```

```python

```